<a href="https://colab.research.google.com/github/talibilat/ce880labcoursework/blob/main/talibilat/ce880labcoursework/week2/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and return log2 of the array. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 4, 8]` and and output should be:  `[0, 2, 3]` 

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myLog(mylist):
    """Write a function to pass a list of numbers and return log2 of the array. """
    z = np.array(mylist)
    return z
    raise NotImplementedError()
```

```python
in_array = [1, 4, 8]
myLog(in_array)
```

Output:
```
array([1, 4, 8])
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case 1
in_array = [1, 3, 5, 10**8]
mynplist = np.array([0, 1.5849625, 2.32192809 , 26.57542476])
np.allclose(myLog(in_array), mynplist)

# Test case 2
in_array = in_array = [1, 100, 1000]
mynplist = np.array([0., 6.64385619, 9.96578428])
np.allclose(myLog(in_array), mynplist)
```

Output:
```
False
```

---

## Question 2:

The Fibonacci numbers are the numbers in the following integer sequence.
`[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, ...]`
In mathematical terms, the sequence $F_n$ of Fibonacci numbers is defined by the recurrence relation 
$$
F_n = F_n-1 + F_n-2
$$
with seed values  
$$
 F_0 = 0 \& F_1 = 1.
$$
Your task here is to write a code for printing $n^{th}$ Fibonacci number
`Example`: Input: `myfibonacci(9)` and Output should be `34`

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myfibonacci(n):
    """Function for nth Fibonacci number"""
    # YOUR CODE HERE
    if n == 0:
      return 0
    elif n == 1:
      return 1
    else:
      number1 = 0
      number2 = 1
      for i in range(2,n+1):
         fib_num = number1 + number2
         number1 = number2
         number2 = fib_num
    
    return (fib_num)

    raise NotImplementedError()
```

```python
myfibonacci(6)
```

Output:
```
8
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct fibonacci number"""
assert myfibonacci(9) == 34
assert myfibonacci(7) == 13
assert myfibonacci(1) == 1
assert myfibonacci(0) == 0
```

---
## Question 3: 

 Write a program, which accept 2D array with `nan` and replace all `nan` with `-5`.  

<!-- grade_id: rem_Nan_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def rem_Nan(A):
    """Write a program, which accept 2D array with nan and replace all nan with -10"""
    
    # YOUR CODE HERE
    A[np.isnan(A)] = -5
    return A
    raise NotImplementedError()
```

```python
arr1 = np.array([[1,2,3],[3,np.nan,3]])
arr2 = np.array([[1,2,3],[3,-5,3]])

assert np.allclose(rem_Nan(arr1), arr2)
```

<!-- grade_id: rem_Nan_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test Case 1
arr1 = np.array([[1,5,3],[3,np.nan,3]])
arr2 = np.array([[1,5,3],[3,-5,3]])
assert np.allclose(rem_Nan(arr1), arr2)

# Test Case 2
arr3 = np.array([[np.nan,2,np.nan],[3,np.nan,3]])
arr4 = np.array([[-5,2,-5],[3,-5,3]])
assert np.allclose(rem_Nan(arr3), arr4)
```

## Pandas Data Frame 


```python
# Please make sure if pandas is installed and if not then run the this cell otherwise skip
!pip install pandas
```

#### Let's create a Toy Data for rest of the 3 questions

```python
import numpy as np
import pandas as pd
exam_data  = {'name': ['Adam', 'Jo', 'Jack', 'Hugo', 'Amit', 'Messi', 'Jordan', 'Laura', 'Kevin', 'Jonas'],
        'score': [72.5, 59, 66.5, np.nan, 59, 80, 84.5, np.nan, 82, 49],
        'attempts': [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
        'qualify': ['yes', 'no', 'yes', 'no', 'no', 'yes', 'yes', 'no', 'no', 'yes']}
labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
df = pd.DataFrame(exam_data , index=labels)
df
```

Output:
```
name  score  attempts qualify
a    Adam   72.5         1     yes
b      Jo   59.0         3      no
c    Jack   66.5         2     yes
d    Hugo    NaN         3      no
e    Amit   59.0         2      no
f   Messi   80.0         3     yes
g  Jordan   84.5         1     yes
h   Laura    NaN         1      no
i   Kevin   82.0         2      no
j   Jonas   49.0         1     yes
```

---
## Question 4:

Write a Pandas program to display students names, where number of attempts in the examination is greater than 2 and score less than 80.

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def filter_1():
    """Write a Pandas program to display students names, where number of attempts 
    in the examination is greater than 2 and score less than 80. c v """
    
    # YOUR CODE HERE
    df1 = df[(df['attempts']) > 2 & (df['score'] < 80)]
    return df1
    raise NotImplementedError()
```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert filter_1() == ['Jo']
```

Output:
```
/usr/local/lib/python3.7/dist-packages/pandas/core/ops/__init__.py in align_method_FRAME(left, right, axis, flex, level)
    279             )
    280         # GH17901
--> 281         right = to_series(right)
    282 
    283     if flex is not None and isinstance(right, ABCDataFrame):

/usr/local/lib/python3.7/dist-packages/pandas/core/ops/__init__.py in to_series(right)
    237             if len(left.columns) != len(right):
    238                 raise ValueError(
--> 239                     msg.format(req_len=len(left.columns), given_len=len(right))
    240                 )
    241             right = left._constructor_sliced(right, index=left.columns)

ValueError: Unable to coerce to Series, length must be 4: given 1
```

## Question 5

Write a Pandas program to display the average score. `Hint`: Please replace `NaN` values with `0` before calcualting average. 

<!-- grade_id: score_ans, points: 0, grade: false, solution: true -->
```python
import itertools
def score():
    """Write a Python program to calculate average score and replace NaN values by 0"""
    # YOUR CODE HERE
    df1 = df[np.isnan(df)] = 0
    return df1.mean
    
    raise NotImplementedError()
```

<!-- grade_id: score_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert score() == 55.250000
```

Output:
```
/usr/local/lib/python3.7/dist-packages/pandas/core/internals/managers.py in apply(self, f, align_keys, ignore_failures, **kwargs)
    323             try:
    324                 if callable(f):
--> 325                     applied = b.apply(f, **kwargs)
    326                 else:
    327                     applied = getattr(b, f)(**kwargs)

/usr/local/lib/python3.7/dist-packages/pandas/core/internals/blocks.py in apply(self, func, **kwargs)
    379         """
    380         with np.errstate(all="ignore"):
--> 381             result = func(self.values, **kwargs)
    382 
    383         return self._split_op_result(result)

TypeError: ufunc 'isnan' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

---
## Question 6:


Write a function to calculate the `mean` and `standard deviation` of the examination `attempts`. `Hint`: You can use NumPy aggregate functions and return as a values. Create a list and append value to it and return the output.

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats():
    """Write a Pandas program to calculate `mean` and `standard deviation` of the examination
    attempts by the students."""
    # YOUR CODE HERE
     
    raise NotImplementedError()
```

<!-- grade_id: stats_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
output = np.array([1.9, 0.8755950357709131])
np.allclose(stats(), output)

```